Before you begin, execute this cell to import numpy and packages from the D-Wave Ocean suite, and all necessary functions for the gate-model framework you are going to use, whether that is the Forest SDK or Qiskit. In the case of Forest SDK, it also starts the qvm and quilc servers.

```python
%run -i "assignment_helper.py"
%matplotlib inline
```

Output:
```
/Users/gianlucafugante/anaconda3/lib/python3.8/site-packages/qiskit/aqua/__init__.py:86: DeprecationWarning: The package qiskit.aqua is deprecated. It was moved/refactored to qiskit-terra For more information see <https://github.com/Qiskit/qiskit-aqua/blob/main/README.md#migration-guide>
  warn_package('aqua', 'qiskit-terra')
Available frameworks:
Qiskit
D-Wave Ocean
```

# Quantum Fourier Transform

**Exercise 1** (3 points). The quantum Fourier transform acts on an amplitude-encoded variant of a normalized vector, the quantum state $|x\rangle=\sum_{i=0}^{N-1} x_i |i \rangle$. It maps this state to $|y\rangle=\sum_{k=0}^{N-1} y_k |k \rangle$, where $y_k = \frac{1}{\sqrt{N}}\sum_{j=0}^{N-1}x_je^{\boldsymbol{2\pi i} \frac{jk}{N}}$. If you have a two-dimensional vector to transform, you only need a single qubit, and the transformation reduces to $U |x \rangle = \frac{1}{\sqrt{2}} \left(|0\rangle + e^{\boldsymbol{2\pi i} 0.x_1} |1\rangle\right)$. Implement this circuit. Place your solution in an object called `qft`. 

```python
###
### YOUR CODE HERE
###

qr = QuantumRegister(1)
cr = ClassicalRegister(1)
qft = QuantumCircuit(qr, cr)

qft.h(qr[0])

qft.draw()
```

Output:
```
┌───┐
q2_0: ┤ H ├
      └───┘
c2: 1/═════
```

<!-- grade_id: exercise1, points: 3, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
amplitudes = get_amplitudes(qft)
amplitudes
```

Output:
```
array([0.70710678+0.j, 0.70710678+0.j])
```

**Exercise 2** (1 point). Implement the inverse quantum Fourier transformation on a single qubit. Place your solution in an object called `iqft`. 

```python
###
### YOUR CODE HERE
###
iqr = QuantumRegister(1)
icr = ClassicalRegister(1)
iqft = QuantumCircuit(iqr, icr)

iqft.h(iqr[0])

iqft.draw()
```

Output:
```
┌───┐
q14_0: ┤ H ├
       └───┘
 c4: 1/═════
```

<!-- grade_id: exercise2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
iamplitudes = get_amplitudes(iqft)
iamplitudes
```

Output:
```
array([0.70710678+0.j, 0.70710678+0.j])
```

# Quantum phase estimation

**Exercise 3** (3 points). Phase estimation pivots on the controlled application of powers of a unitary matrix, followed by an inverse Fourier transformation. Take the $2\times 2$ unitary matrix $\begin{bmatrix}1 & 0 \\0 & i\end{bmatrix}$. This operation has two eigenvectors, $|0\rangle$ and $|1\rangle$. Implement the first part of the quantum phase estimation with two ancilla qubits: add the necessary Hadamard gates on the ancilla qubits and apply $C-U^{2^0}$ and $C-U^{2^1}$. Qubits 0 and 1 should be the ancilla, and qubit 2 contains the eigenvector -- you can assume that this is $|0\rangle$. Place your solution in an object called `circuit`. 

```python
###
### YOUR CODE HERE
###

qr = QuantumRegister(3)
cr = ClassicalRegister(3)
circuit = QuantumCircuit(qr, cr)

circuit.h(qr[0])
circuit.h(qr[1])

#circuit.cu1(np.pi/2., qr[0], qr[2])
circuit.cu1(np.pi/2., qr[1], qr[2])
circuit.cz(qr[0], qr[2])

circuit.draw()

```

Output:
```
┌───┐             
q42_0: ┤ H ├───────────■─
       ├───┤           │ 
q42_1: ┤ H ├─■─────────┼─
       └───┘ │U1(π/2)  │ 
q42_2: ──────■─────────■─
                         
 c8: 3/══════════════════
```

<!-- grade_id: exercise3, points: 3, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([0.5+0.j, 0.5+0.j, 0.5+0.j, 0.5+0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j]))
```

**Exercise 4** (1 point). Apply the same circuit to the other eigenvector. Place your solution in an object called `circuit_2`.

```python
###
### YOUR CODE HERE
###
q = QuantumRegister(3)
c = ClassicalRegister(3)
circuit_2 = QuantumCircuit(q, c)
circuit_2.x(q[2])
circuit_2.h(q[0])
circuit_2.h(q[1])
circuit_2.cu1(np.pi/2, q[1], q[2])
circuit_2.cz(q[0], q[2])
```

Output:
```
<qiskit.circuit.instructionset.InstructionSet at 0x12f462c40>
```

<!-- grade_id: exercise4, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit_2)
assert np.allclose(amplitudes, np.array([ 0. +0.j ,  0. +0.j ,  0. +0.j ,  0. +0.j ,  0.5+0.j , -0.5+0.j ,
                                          0. +0.5j,  0. -0.5j]))
```

```python

```